## Before submitting
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

2. Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.

3. After you submit the notebook more tests will be run on your code. The fact that no assertions fail on your computer localy does not guarantee that completed the exercise correctly.

4. Please submit only the `*.ipynb` file.

5. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Edit only between `YOUR CODE HERE` and `END YOUR CODE`.

6. Make sure to use Python 3, not Python 2.

Fill your group name and collaborators below:

```python
GROUPNAME = "group100"
COLLABORATORS = "1.Nautiyal, Amit ; 2.Rech, Florian ; 3.Stevens, Kashmir Chantal"
```

---

# Exercise Sheet 1: Python Basics

This first  exercise sheet tests the basic functionalities of the Python programming language in the context of a simple prediction task. We consider the problem of predicting health risk of subjects from personal data and habits. We first use for this task a decision tree

![](tree.png)

adapted from the webpage http://www.refactorthis.net/post/2013/04/10/Machine-Learning-tutorial-How-to-create-a-decision-tree-in-RapidMiner-using-the-Titanic-passenger-data-set.aspx. For this exercise sheet, you are required to use only pure Python, and to not import any module, including numpy. In exercise sheet 2, the nearest neighbor part of this exercise sheet will be revisited with numpy.

## Classifying a single instance (15 P)

* In this sheet we will represent patient info as a tuple.
* Create a function that takes as input a tuple containing values for attributes (smoker,age,diet), and computes the output of the decision tree. Should return `"less"` or `"more"`.
* Test your function on the tuple `('yes', 31, 'good')`.

<!-- grade_id: cell-b66d7278bc313c94, points: 0, grade: false, solution: true -->
```python
def decision(x):
    if (x[0]=='yes'):
        if(x[1]<29.5):
         return 'less'
        else:
         return 'more'
    else:
        if(x[2]=='good'):
         return 'less'
        else:
         return 'more'
    
    ##dev test 1 
    
test_input = ('yes', 31, 'good')
test = decision(test_input)
assert test == 'more'
answer = ((test_input, test))
answer




```

Output:
```
(('yes', 31, 'good'), 'more')
```

<!-- grade_id: cell-c31b80471db3132f, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert decision(x) == 'more'

```

<!-- grade_id: cell-a706e6de3303d3e7, points: 0, grade: true, solution: false -->
```python

```

## Reading a dataset from a text file (10 P)

The file `health-test.txt` contains several fictious records of personal data and habits.

* Read the file automatically using the methods introduced during the lecture.
* Represent the dataset as a list of tuples. Make sure that the tuples have the same format as above, e.g. `('yes', 31, 'good')`.
* Make sure that you close the file after you have opened it and read its content.

**Notes**: 
* Values read from files are always strings.
* Each line contains a newline `\n` character at the end

<!-- grade_id: cell-c1a8bc4c0e4ccb26, points: 0, grade: false, solution: true -->
```python
input_file ='health-test.txt'
ht = []
with open(input_file,'r') as f:
    for l  in f:
        formatting = l.strip().split(',')
        formatting[1]= int(formatting[1])
        ht.append(tuple(formatting))
##calling        
ht

        
        
   
```

Output:
```
[('yes', 21, 'poor'),
 ('no', 50, 'good'),
 ('no', 23, 'good'),
 ('yes', 45, 'poor'),
 ('yes', 51, 'good'),
 ('no', 60, 'good'),
 ('no', 15, 'poor'),
 ('no', 18, 'good')]
```

<!-- grade_id: cell-4e1f7ad1e66b3121, points: 5, grade: true, solution: false -->
```python
parsed_line = parse_line_test('yes,23,good\n')
assert isinstance(parsed_line, tuple)
assert isinstance(parsed_line[1], int)

```

<!-- grade_id: cell-0ed341eaf90d76a6, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ab1b01fcaac26cd5, points: 5, grade: true, solution: false -->
```python
testset = gettest()
assert isinstance(testset, list)
assert isinstance(testset[0], tuple)
```

<!-- grade_id: cell-42b0a5cad68fc15a, points: 0, grade: true, solution: false -->
```python

```

## Applying the decision tree to the dataset (15 P)

* Apply the decision tree to all points in the dataset, and return the ratio of them that are classified as "more".
* A ratio is a value in [0-1]. So if out of 50 data points 15 return `"more"` the value that should be returned is `0.3`

<!-- grade_id: cell-6703ef98e2b5c93b, points: 0, grade: false, solution: true -->
```python
decision_tree = []
for i in ht:
    decision_tree.append(decision(i))

find_more = float(decision_tree.count('more')) / len(decision_tree)
print("the dataset is %s:"%decision_tree)
print("RATIO: %s" % find_more)

#call function
find_more
    
```

Output:
```
the dataset is ['less', 'less', 'less', 'more', 'more', 'less', 'more', 'less']:
RATIO: 0.375
0.375
```

<!-- grade_id: cell-c13a0b23c9faba52, points: 15, grade: true, solution: false -->
```python
ratio = evaluate_testset()
assert isinstance(ratio, float)
assert 0.0 <= ratio <= 1.0

```

## Learning from examples (10 P)

Suppose that instead of relying on a fixed decision tree, we would like to use a data-driven approach where data points are classified based on a set of training observations manually labeled by experts. Such labeled dataset is available in the file `health-train.txt`. The first three columns have the same meaning than for `health-test.txt`, and the last column corresponds to the labels.

* Write a function that reads this file and converts it into a list of pairs. The first element of each pair is a triplet of attributes, and the second element is the label.

**Note**: A triplet is a tuple that contains exactly three values, a pair is a tuple that contains exactly two values

<!-- grade_id: cell-fc38ed11fee6fbeb, points: 0, grade: false, solution: true -->
```python
input_file ='health-train.txt'
hlthtrn = []
with open(input_file,'r') as fl:
   
    print('---------------------the tuple is------------------------ ')

    for li  in fl:
        new_formatting = li.strip().split(',')
        new_formatting[1]= int(new_formatting[1])
        #formating :the first element of each pair is a triplet of attributes, and the second element is the label
        hlthtrn.append((tuple(new_formatting[:3]),new_formatting[-1]))
        
##calling        
hlthtrn

  
    
```

Output:
```
---------------------the tuple is------------------------ 
[(('yes', 54, 'good'), 'less'),
 (('no', 55, 'good'), 'less'),
 (('no', 26, 'good'), 'less'),
 (('yes', 40, 'good'), 'more'),
 (('yes', 25, 'poor'), 'less'),
 (('no', 13, 'poor'), 'more'),
 (('no', 15, 'good'), 'less'),
 (('no', 50, 'poor'), 'more'),
 (('yes', 33, 'good'), 'more'),
 (('no', 35, 'good'), 'less'),
 (('no', 41, 'good'), 'less'),
 (('yes', 30, 'poor'), 'more'),
 (('no', 39, 'poor'), 'more'),
 (('no', 20, 'good'), 'less'),
 (('yes', 18, 'poor'), 'less'),
 (('yes', 55, 'good'), 'more')]
```

<!-- grade_id: cell-b97620858b167c0d, points: 5, grade: true, solution: false -->
```python
parsed_line = parse_line_train('yes,67,poor,more\n')
assert isinstance(parsed_line, tuple)
assert isinstance(parsed_line[0], tuple)
assert isinstance(parsed_line[1], str)
assert parsed_line[0][1] == 67

```

<!-- grade_id: cell-ced9925599b5dd99, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a3d593f232e0403a, points: 5, grade: true, solution: false -->
```python
trainset = gettrain()
assert isinstance(trainset, list)
assert isinstance(trainset[0], tuple)
assert isinstance(trainset[0][0], tuple)
assert isinstance(trainset[0][1], str)
```

<!-- grade_id: cell-7ab8ef2c931d69c1, points: 0, grade: true, solution: false -->
```python

```

## Nearest neighbor classifier (25 P)

We consider the nearest neighbor algorithm that classifies test points following the label of the nearest neighbor in the training data. For this, we need to define a distance function between data points. We define it to be

`distance(a, b) = (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])`

where `a` and `b` are two tuples corrsponding to the attributes of two data points.

* Write a function that retrieves for a test point the nearest neighbor in the training set, and classifies the test point accordingly (i.e. returns the label of the nearest data point).
* Test your function on the tuple `('yes', 31, 'good')`

**Note**: You can use the special `infinity` floating point value with `float('inf')`

<!-- grade_id: cell-671ea24ec8a11241, points: 0, grade: false, solution: true -->
```python
def distance(a,b):
     return (a[0] != b[0]) + (a[0]!=b[0]) + ((float(a[1])-float(b[1]))/50)**2 + (a[2]!=b[2])

def neighbor(x, training_set):
    ##claculation in traing set 
    dist = [distance(x, y[0]) for y in training_set]
    #return it 
    return training_set[dist.index(min(dist))]


    
```

<!-- grade_id: cell-5929fce6ffc8ca0f, points: 5, grade: true, solution: false -->
```python
# Test distance
import math
assert math.isclose(distance(('yes', 34, 'poor'), ('yes', 51, 'good')), 1.1156)

```

<!-- grade_id: cell-8dc0c32bc78d1e62, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a36122337853f195, points: 10, grade: true, solution: false -->
```python
# Test neighbor
x = ('yes', 31, 'good')
assert neighbor(x, gettrain()) == 'more'

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-3f38acbf2846> in <module>
      1 # Test neighbor
      2 x = ('yes', 31, 'good')
----> 3 assert neighbor(x, gettrain()) == 'more'

NameError: name 'gettrain' is not defined
```

<!-- grade_id: cell-cbf6ddfbfaab7875, points: 0, grade: true, solution: false -->
```python

```

* Apply both the decision tree and nearest neighbor classifiers on the test set, and return the list of data point(s) for which the two classifiers disagree, and with which probability it happens.
* A data point should look like above, e.g. `('yes', 31, 'good')`.

<!-- grade_id: cell-8dbf7da153f3d797, points: 0, grade: false, solution: true -->
```python
health_nearest = []

index = []
diff = []

for i in ht:
    formatting = neighbor(i, hlthtrn)
    health_nearest.append(formatting[1])

print(" Decision Tree Prediction:\n%s" % decision_tree)

print(" Nearest Neighbor Algorithm Predictions :\n%s" % health_nearest)


for indices in range(len(decision_tree)):
    if decision_tree[indices] != health_nearest[indices]:
        index.append(indices)
        diff.append(ht[indices])

ProbablityOfDiff = float(len(diff)) / len(decision_tree)

print("Index: %s\nDatapoint: %s" % (index, diff))
print("Probability: %s" % ProbablityOfDiff)
print()


((diff, ProbablityOfDiff))
```

Output:
```
Decision Tree Prediction:
['less', 'less', 'less', 'more', 'more', 'less', 'more', 'less']
 Nearest Neighbor Algorithm Predictions :
['less', 'less', 'less', 'more', 'less', 'less', 'more', 'less']
Index: [4]
Datapoint: [('yes', 51, 'good')]
Probability: 0.125

([('yes', 51, 'good')], 0.125)
```

<!-- grade_id: cell-3b55f7e89ad4dfeb, points: 10, grade: true, solution: false -->
```python
Xdisagree, probability = compare_classifiers()
assert isinstance(Xdisagree, list)
assert 0.0 <= probability <= 1.0
```

One problem of simple nearest neighbors is that one needs to compare the point to predict to all data points in the training set. This can be slow for datasets of thousands of points or more. Alternatively, some classifiers train a model first, and then use it to classify the data.

## Nearest mean classifier (25 P)

We consider one such trainable model, which operates in two steps:

(1) Compute the average point for each class, (2) classify new points to be of the class whose average point is nearest to the point to predict.

For this classifier, we convert the attributes smoker and diet to real values (for smoker: yes=1.0 and no=0.0, and for diet: good=0.0 and poor=1.0), and use the modified distance function:

`distance(a,b) = (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2`

Age will also from now on be represented as a `float`. The new data points will be referred to as numerical patient tuples. 

We adopt an object-oriented approach for building this classifier.

* Implement the methods `train` and `predict` of the class `NearestMeanClassifier`.

<!-- grade_id: cell-e0b339bfd0fcc16c, points: 0, grade: false, solution: true -->
```python
class NearestMeanClassifier(object):
   
    def train(self, dataset):
        
        self.classes = []
        data = []
        self.avg = []
        
        for l in dataset:
            if l[1] not in self.classes:
                self.classes.append(l[1])

            X, y = l
            smoker, age, diet = X
            tuple1 = ((int(smoker == 'yes'),
                              int(age),
                              int(diet == 'poor')), y)
            data.append(tuple1)

        

        for classes in self.classes:
            
            listOfsmoker = [tuples[0][0] for tuples in data if tuples[1] == classes]
            age = [tuples[0][1] for tuples in data if tuples[1] == classes]
            diet = [tuples[0][2] for tuples in data if tuples[1] == classes]
            dict_content = [tuples for tuples in data if tuples[1] == classes]

            meanOfAbove = ((float(sum(listOfsmoker)) / len(listOfsmoker),
                           float(sum(age)) / len(age),
                           float(sum(diet)) / len(diet)), classes)
            self.avg.append(meanOfAbove)

    def predict(self, dataset):
        
        ##as defined earlier
        def distance(a, b):
            return (a[0] - b[0])**2 + ((a[1] - b[1])/50)**2 + (a[2] - b[2])**2

        def neighbor(x, training_set):
            distbtw = [distance(x, y[0]) for y in training_set]
            return training_set[distbtw.index(min(distbtw))][1]


        self.predict = []
        for l2 in dataset:
            smoker, age, diet = l2
            tuple_line = (int(smoker == 'yes'),
                            int(age),
                            int(diet == 'poor'))

            prcdt = neighbor(tuple_line, self.avg)
            self.predict.append(prcdt)

```

* Implement a function that will load the training dataset from the `health-train.txt` file and parse each line to a numerical patient tuple. You can still follow the same structure that we used before (i.e. using a `parse_line` function), however, it is not required for this exercise. 
* Build an object of class `NearestMeanClassifier`, train it on the training data, and return the mean class numerical tuple for each class.

<!-- grade_id: cell-5f7f00ee83c94703, points: 0, grade: false, solution: true -->
```python
mean_class = NearestMeanClassifier()

##load the training dataset from the health-train.txt
mean_class.train(hlthtrn)
mean_class.predict(ht)

mean_class.avg
```

Output:
```
[((0.3333333333333333, 32.111111111111114, 0.2222222222222222), 'less'),
 ((0.5714285714285714, 37.142857142857146, 0.5714285714285714), 'more')]
```

<!-- grade_id: cell-0a0dea69eca74de7, points: 5, grade: true, solution: false -->
```python
trainset_num = gettrain_num()
assert isinstance(trainset_num, list)
assert isinstance(trainset_num[0][0][0], float)

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-135-6c8d875bb73a> in <module>
----> 1 trainset_num = gettrain_num()
      2 assert isinstance(trainset_num, list)
      3 assert isinstance(trainset_num[0][0][0], float)

NameError: name 'gettrain_num' is not defined
```

<!-- grade_id: cell-789233fd9f800e77, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-415891bde4cbde19, points: 2, grade: true, solution: false -->
```python
train_dict = build_and_train()
assert isinstance(train_dict['less'], tuple)
assert isinstance(train_dict['more'], tuple)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-136-70ccf2a0e2d7> in <module>
----> 1 train_dict = build_and_train()
      2 assert isinstance(train_dict['less'], tuple)
      3 assert isinstance(train_dict['more'], tuple)

NameError: name 'build_and_train' is not defined
```

<!-- grade_id: cell-ca2d0921e96ede25, points: 8, grade: true, solution: false -->
```python

```

* Load the test dataset into memory as a list of numerical patient tuples
* Predict the test data using the nearest mean classifier and print all test examples for which all three classifiers (decision tree, nearest neighbor and nearest mean) agree.

**Note**: Be careful that the `NearestMeanClassifier` expects the dataset in a different form, compared to the other two methods.

<!-- grade_id: cell-f37f3035a32a8f85, points: 0, grade: false, solution: true -->
```python
compare = []
for i in range(len(ht)):
    
    if decision_tree[i] == health_nearest[i]:
        compare.append(ht[i])
    
print("the similar  Predictions are :")
compare
```

Output:
```
the similar  Predictions are :
[('yes', 21, 'poor'),
 ('no', 50, 'good'),
 ('no', 23, 'good'),
 ('yes', 45, 'poor'),
 ('no', 60, 'good'),
 ('no', 15, 'poor'),
 ('no', 18, 'good')]
```

<!-- grade_id: cell-f3656461c994dc3d, points: 2, grade: true, solution: false -->
```python
testset_num = gettest_num()
assert isinstance(testset_num, list)
assert isinstance(testset_num[0], tuple)
```

<!-- grade_id: cell-853c957eaaf81c28, points: 8, grade: true, solution: false -->
```python
same_predictions = predict_test()
assert isinstance(same_predictions[0], tuple)
assert isinstance(same_predictions[0][0], tuple)
assert isinstance(same_predictions[0][1], str)
```